# Lab 1: Data Types, Functions and Conditionals (29 points)

#### Exercise 0
Execute the code and try to understand the error msg

```python
print('Hello World') #incomplete syntax, missing paranthesis
```

```python
'Welcome'/2 #cant divide str by int
```

```python
print(a) # a aint declared
```

Output:
```
NameError                                 Traceback (most recent call last)
/tmp/ipykernel_5676/360640881.py in <cell line: 0>()
----> 1 print(a) # a aint declared

NameError: name 'a' is not defined
```

#### Exercise 1
 Write a function which computes gross pay given the **hours** worked and the **rate_per_hour**.


<!-- grade_id: cell-312caeaaf6abe1a9, points: 0, grade: false, solution: true -->
```python
def get_gross_pay(hours, rate_per_hour):
    # YOUR CODE HERE
    gross_pay = hours*rate_per_hour
    print(gross_pay)
    #raise NotImplementedError()
hr = int(input("Enter hours worked: "))
rph = int(input("Enter rate per hour: "))
get_gross_pay(hr, rph)
```

Output:
```
Enter hours worked: 4
Enter rate per hour: 10
40
```

<!-- grade_id: cell-458cb3b5b6333645, points: 1, grade: true, solution: false -->
```python
assert get_gross_pay(3, 10) == 30
```

#### Exercise 2
Assume that we execute the following assignment statements.
```
width = 17
height = 12.0
```
For each of the following expressions, write the value of the expression and the type (of the value of the expression).

1. `width//2`

2. `width/2.0`

3. `height/3`

4. `1 + 2 * 5`

Execute the above code to check.
1. float, 8.0
2. float, 8.5
3. float, 4.0
4. int, 11

#### Exercise 3
Rewrite your `get_gross_pay` function from ch2, exercise 3 to give the employee 1.5 times the hourly rate for hours worked above 40 hours.

<!-- grade_id: cell-8f610a982f866fac, points: 0, grade: false, solution: true -->
```python
def get_gross_pay(hours, rate_per_hour):
    # YOUR CODE HERE
    gross_pay = hours*rate_per_hour
    return gross_pay
    raise NotImplementedError()
get_gross_pay(40, 1.5)
```

Output:
```
60.0
```

<!-- grade_id: cell-537d9d1c5dcd290f, points: 1, grade: true, solution: false -->
```python
assert get_gross_pay(0, 1000)==0
assert get_gross_pay(2, 10)==20
```

#### Exercise 4
What is the purpose of the `def` keyword in Python?

a) It is slang that means “the following code is really cool”

b) It indicates the start of a function

c) It indicates that the following indented section of code is to be stored for later

#d) b and c are both true

e) None of the above

<!-- grade_id: ch4-ex1-ans, points: 0, grade: false, solution: true -->
```python
# Enter your answer as a python string in a variable 'ans'.
# Should be one of ['a', 'b', 'c', 'd', 'e']
# YOUR CODE HERE
ans = 'd'
#raise NotImplementedError()
```

<!-- grade_id: ch4-ex1-test, points: 1, grade: true, solution: false -->
```python
assert type(ans) == str
assert ans in ['a','b','c','d','e']
```

#### Exercise 5
Create a new GitHub repository for the Data Science course and complete the following tasks:
### Set up the repository
Name the repository appropriately (e.g., "stats507-coursework");
Ensure the repository is public (unless otherwise specified)

### Add a README file
Create a file named "README.md" in the root directory;
Include a brief description of the course and the purpose of the repository

### Provide the repository link


<!-- grade_id: cell-9a484821ec48c9fc, points: 0, grade: false, solution: true -->
```python
# YOUR REPO LINK HERE
raise NotImplementedError()
```